---
title       : Insert the chapter title here
description : Insert the chapter description here
attachments :
  slides_link : https://s3.amazonaws.com/assets.datacamp.com/course/teach/slides_example.pdf

---
## Milline dplyr versioon siin on?

```yaml
type: NormalExercise
key: 3fa84fa18e
lang: r
xp: 100
skills: 1
```


`@instructions`
Tee midagi

`@hint`

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}
library(dplyr)


sessionInfo()

```

`@solution`
```{r}

```

`@sct`
```{r}

```
